# W8 PIT Practical 

# Fibonacci number
---

The Fibonacci numbers are `0, 1, 1, 2, 3, 5, 8, 13, 21...`. 

The formula is `Fn = Fn-1 + Fn-2`, with `F0=0` and `F1=1`. 

For example,
- `F2 = F1 + F0 = 1 + 0 = 1`
-  `F3 = F2 + F1 = 1 + 1 = 2`
-  `F4 = F3 + F2 = 2 + 1 = 3`
- `F5 = F4 + F3 = 3 + 2 = 5`
- ...

## Task 1 (25%)
Please complete the function `fibo` below that displays Fibonacci number using a **non-recursive** approach

```python
def fibo(n):
    i = 1
    f0 = 0
    f1 = 1
    fi = 1
    while i < n+1:
        print(f'f{i} = {fi}')    
        # INPUT YOUR CODE BELOW
        # approx 3 lines
        ...
        fi = f1 + f0
        f0 = f1
        f1 = fi
        i += 1
        
# =======================
# EXAMPLE TESTING CASES
fibo(5)
# output:
# f1 = 1
# f2 = 1
# f3 = 2
# f4 = 3
# f5 = 5
```

Output:
```
f1 = 1
f2 = 1
f3 = 2
f4 = 3
f5 = 5
```

## Task 2 (25%)
Please display Fibonacci number using a recursive approach

```python
def fibo_rec(n):
    # INPUT YOUR CODE BELOW
    # approx 4 lines
    ...
    # base case     
    if n==0: 
        return 0     
    if n==1: 
        return 1     
    
    # recursive case     
    return fibo_rec(n-2)+fibo_rec(n-1) 
    
# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
n = 10
print(f'f{n} is: {fibo_rec(n)}') # f10 is: 55

n = 20
print(f'f{n} is: {fibo_rec(n)}') # f20 is: 6765
```

Output:
```
f10 is: 55
f20 is: 6765
```

## Task 3 (25%)
Please implement the quick sort algorithm that 
1. sort the numbers in **ascending** order and
1. uses the **first** element as the pivot.

```python
import random

def qsort(arr):
    # YOUR CODE BELOW
    # ~ 2 lines, the base case    
    ...
    if(len(arr)==0):         
        return arr
    else:
        # Read two lines below carefully, you will implement the partition function
        lesser, greater = partition(arr)
        return qsort(lesser) + [arr[0]] + qsort(greater) 
    
def partition(arr):
    lesser = []
    greater = []
    pivot = arr[0]
    # YOUR CODE BELOW
    # ~ 5 lines, create two lists: lesser and greater
    ...
    for index in range(len(arr)):         
        if(arr[index]<pivot):             
            lesser.append(arr[index])        
        elif(arr[index]>pivot):             
            greater.append(arr[index])
    return lesser, greater
    
# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
A = [12, 23, 43, 21]
arr = qsort(A)
print(arr) #[12, 21, 23, 43]
```

Output:
```
[12, 21, 23, 43]
```

## Exercise 4 (25%)
Please implement the quick sort algorithm that 
1. sort the numbers in **descending** order and
1. uses **random** element as the pivot.

For simplicity, you can assume all number in the array are unique.

```python
import random

def qsort_dec(arr):
    # YOUR CODE BELOW
    # ~5 LINES
    # Copy and modify the qsort() function above
    if(len(arr)==0):
        return arr
    else:
        lesser, greater, pivot = partition_dec(arr)
        return qsort_dec(greater) + [pivot] + qsort_dec(lesser)
    
def partition_dec(arr):
    lesser = []
    greater = []
    # YOUR CODE BELOW    
    # ~ 6 lines, copy and update from Task 3
    pivot = random.choice(arr)
    ...
    for index in range(len(arr)):         
        if(arr[index]<pivot):             
            lesser.append(arr[index])        
        elif(arr[index]>pivot):             
            greater.append(arr[index])
    
    return lesser, greater, pivot 

# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
A = [5, 9, 12, 16, 23]
arr = qsort_dec(A)
print(arr) # [23, 16, 12, 9, 5]
```

Output:
```
[23, 16, 12, 9, 5]
```

```python

```

Tests for Task 1 (points: 1)
```
>>> fibo(10)
f1 = 1
f2 = 1
f3 = 2
f4 = 3
f5 = 5
f6 = 8
f7 = 13
f8 = 21
f9 = 34
f10 = 55
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task 2 (points: 1)
```
>>> fibo_rec(5) == 5
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> fibo_rec(10) == 55
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task 3 (points: 1)
```
>>> a = qsort([5, 3, 2, -5])
>>> a[3] ==5 
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task 4 (points: 1)
```
>>> qsort_dec([98, 23, 43, 21, 53, 29])
[98, 53, 43, 29, 23, 21]
```
failure_message: Incorrect answer, try again!
success_message: Good job!
